<a href="https://colab.research.google.com/github/iammrmahid95/CE880_Lab_coursework/blob/main/week7/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mohammed Mahbubur Rob"
COLLABORATORS = "Mohammed Mahbubur Rob"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

---
## Question 1:
#### Bootstrapping
We are creating a random vector to serve as our population, with an unknown distribution matching that of the population. Develop a function to select a sample from our population using the Bootstrapping procedure, with a sample size of `300` and conducting `1000` re-samplings (`B = 1000`).

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=1000)
print(pop.mean()) # Print population mean

sample = np.random.choice(pop, size=300) #so n=300 #example of picking up a sample

def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    sample_means = []

    for _ in range(1000):
        # Sample with replacement
        sample = np.random.choice(pop, size=300, replace=True)
        # Calculate the mean of the sample and store it
        sample_mean = np.mean(sample)
        sample_means.append(sample_mean)

    # Return the mean of the sample means
    return np.mean(sample_means)
    raise NotImplementedError()
```

Output:
```
253.241
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

---
## Question 2:

#### Majority Voting
Given a list = `[1,2,3,2,2,1,3,3,3,2,2,1]`, consider 12 judges sitting and deciding among 3 contenders. Write a function to find the winner-based on a majority voting scheme.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
a = np.array([1,2,3,2,2,1,3,3,3,2,2,1])

def my_majorityvote():
    """Write a function to find the majority voting."""
    unique_contenders, counts = np.unique(a, return_counts=True)
    winner = unique_contenders[np.argmax(counts)]
    return winner
    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 2, rel_tol = 0.05)
```

---
## Question 3:

#### Decision Tree
USe the IRIS dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely IRIS Setosa, IRIS Versicolour, and IRIS Virginica, with the following attributes: sepal length, sepal width, petal length, and petal width. Your aim is to train a decision tree on training data and return the test accuracy.

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    # Create a decision tree classifier
    dt_classifier = DecisionTreeClassifier(random_state=42)

    # Train the classifier on the training data
    dt_classifier.fit(X_train, y_train)

    # Predict the labels for the test set
    y_pred = dt_classifier.predict(X_test)

    # Calculate and return the test accuracy
    accuracy = accuracy_score(y_test, y_pred)
    return accuracy
    raise NotImplementedError()
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 0.94736, rel_tol = 0.05)
```