<a href="https://colab.research.google.com/github/iammrmahid95/CE880_Lab_coursework/blob/main/week6/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mohammed Mahbubur Rob"
COLLABORATORS = "Mohammed Mahbubur Rob"
```

---

---
# Welcome to CE880
### This is your week-6 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
#!pip install -U scikit-learn
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1:

#### Student’s t-test
Tests whether the means of two independent samples are significantly different and return p-value. Please consider $\alpha = 0.05$

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_ttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import ttest_ind
data1 = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]
data2 = [1.142, -0.432, -0.938, -0.729, -0.846, -0.157, 0.500, 1.183, -1.075, -0.169]

def my_ttest(data1, data2):
    """Write a code to tests whether the means of two independent samples are significantly
    different and return p-value. Please consider 𝛼 = 0.05  """
    # Perform t-test
    t_stat, p_value = ttest_ind(data1, data2)
    return p_value
    raise NotImplementedError()
```

<!-- grade_id: my_ttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_ttest(data1, data2), 0.748, rel_tol = 0.05)
```

---
## Question 2:

#### Paired Student’s t-test
Tests whether the means of two paired samples are significantly different. Please consider $\alpha = 0.05$

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.
* Observations across each sample are paired.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
from scipy.stats import ttest_rel
data1 = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]
data2 = [1.142, -0.432, -0.938, -0.729, -0.846, -0.157, 0.500, 1.183, -1.075, -0.169]


def my_pairedttest(data1, data2):
    """Write code to tests whether the means of two paired samples are
    significantly different. Please consider 𝛼=0.05 """
    t_stat, p_value = ttest_rel(data1, data2)

    return p_value
    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_pairedttest(data1, data2), 0.746, rel_tol = 0.05)
```

---
## Question 3:

#### Shapiro-Wilk Test
Tests whether a data sample has a Gaussian distribution.

Assumptions

* Observations in each sample are independent and identically distributed (iid).

Interpretation

* H0: the sample has a Gaussian distribution.
* H1: the sample does not have a Gaussian distribution.

<!-- grade_id: my_normalityttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import shapiro
data = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]

def my_normalityttest(data):
    """Write code to to check the normality of the data. In other words,
    you are checking if the data is normally distributed and return p value.
    Hint: Please see what is shapiro test and how you can use it
    https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.shapiro.html"""
    stat, p_value = shapiro(data)

    return p_value
    raise NotImplementedError()

my_normalityttest(data)
```

Output:
```
0.1934097558259964
```

<!-- grade_id: my_normalityttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_normalityttest(data), 0.193, rel_tol = 0.05)
```

```python

```