<a href="https://colab.research.google.com/github/iammrmahid95/CE880_Lab_coursework/blob/main/week10/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mohammed Mahbubur Rob"
COLLABORATORS = "Mohammed Mahbubur Rob"
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

Paris Housing is a set of data created from imaginary data of house prices in an urban environment.
https://github.com/sagihaider/CE880_2021/blob/main/Data/ParisHousingClass.csv

All attributes are numeric variables and they are listed bellow:
* squareMeters
* numberOfRooms
* hasYard
* hasPool
* floors - number of floors
* cityCode - zip code
* cityPartRange - the higher the range, the more exclusive the neighbourhood is
* numPrevOwners - number of prevoious owners
* made - year
* isNewBuilt
* hasStormProtector
* basement - basement square meters
* attic - attic square meteres
* garage - garage size
* hasStorageRoom
* hasGuestRoom - number of guest rooms
* price - price of a house
* category - Luxury or Basic

```python
import numpy as np
import pandas as pd
url = 'https://raw.githubusercontent.com/sagihaider/CE880_2021/main/Data/ParisHousingClass.csv'
trainData = pd.read_csv(url,index_col=0)
print(trainData.head())
```

Output:
```
numberOfRooms  hasYard  hasPool  floors  cityCode  \
squareMeters                                                      
75523                     3        0        1      63      9373   
80771                    39        1        1      98     39381   
55712                    58        0        1      19     34457   
32316                    47        0        0       6     27939   
70429                    19        1        1      90     38045   

              cityPartRange  numPrevOwners  made  isNewBuilt  \
squareMeters                                                   
75523                     3              8  2005           0   
80771                     8              6  2015           1   
55712                     6              8  2021           0   
32316                    10              4  2012           0   
70429                     3              7  1990           1   

              hasStormProtector  basement  attic  garage  hasStorageRoom  \
squareMeters       
```

```python
from sklearn.preprocessing import LabelEncoder

def label_encoded(feat):
    le=LabelEncoder()
    le.fit(feat)
    print(feat.name,le.classes_)
    return le.transform(feat)
trainData['category']=label_encoded(trainData['category'])
```

Output:
```
category ['Basic' 'Luxury']
```

```python
from sklearn.model_selection import train_test_split,GridSearchCV

y=trainData['category']
x=trainData.drop('category',axis=1)
xtrain,xtest,ytrain,ytest=train_test_split(x,y,test_size=0.2,random_state=42)
```

---
## Question 1:

#### Classification
Use the dataset, which I have spilted into training and test sets (xtrain, xtest, ytrain, ytest), please train the machine learning model to get the test accuracy of 100%. You are free to use any classication model such as Decision Tree, Random Forest, KNN, SVM..etc.

Hint: Please use grid search to find a best model using different parameters.

<!-- grade_id: my_model_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score
def my_model(xtrain,xtest,ytrain,ytest):
    """Write a function to train a machine learning model to get the test accuracy of 100%.
    You are free to use any classication model such as Decision Tree,
    Random Forest, KNN, SVM..etc. """
    # Define the model
    model = RandomForestClassifier()

    # Define hyperparameters to search
    param_grid = {
        'n_estimators': [50, 100, 150],
        'max_depth': [None, 10, 20],
        'min_samples_split': [2, 5, 10],
        'min_samples_leaf': [1, 2, 4]
    }

    # Perform GridSearchCV
    grid_search = GridSearchCV(model, param_grid, cv=5, scoring='accuracy')
    grid_search.fit(xtrain, ytrain)

    # Get the best model
    best_model = grid_search.best_estimator_

    # Make predictions on the test set
    y_pred = best_model.predict(xtest)

    # Calculate test accuracy
    test_accuracy = accuracy_score(ytest, y_pred)

    return test_accuracy

    raise NotImplementedError()


```

<!-- grade_id: my_model_test, points: 9, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_model(xtrain,xtest,ytrain,ytest), 100.0, rel_tol = 0.05)
```

```python

```